# Spaghetti code
(15 points)

You are given a set of spaghetti in different lengths and the goal is to have as many spaghetti as possible with an exact length of 10cm. Note that 
* it is allowed to cut spaghettis at every position
* it is _not_ possible to attach them together
* there is an overall maximum number of cuts that are allowed to be done

The goal is to find the maximum number of 10cm pieces that can be created based on the given set of spaghettis and the given maximum number of cuts.

For example, if a spaghetti of size 33cm is cutted after the first 10cm, two new pieces are created with 10cm and 23cm length. Following this approach, the `33` can be transformed to `{ 10, 10, 10, 3 }`. Using 3 cuts, 3 spaghettis of length 10 can be created from a spaghetti of size 33 (assuming that the maximum number of cuts is tree or more). Please note that the maximum number of cuts is an overall maximum. So if the maximum was 3 cuts and the implemented algorithm decides to use all three cuts for the 33cm spaghetti, there are no more cuts left for another spaghetti.

Input: `int[] spaghetti, int maximumCuts`

Output: the maximum number of spaghetti of length 10 that you can get by atmost `maximumCut` cuts as `int`

Example Input: `{13,14,20}, 2`

Example Output: `3`

#### Notes

- You are free to use a different IDE to develop your solution. However, you have to copy the solution into this notebook to submit it.
- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Known issues
  - All global variables will be set to void after an import.
  - Missing spaces arround `%` (Modulo) can cause unexpected errors so please make sure that you have added spaces around every `%` character.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Make sure that
    - all necessary imports are listed at the beginning of your cell.
  - Run a final check of your solution by
    - click on _restart the kernel, then re-run the whole notebook_ (the fast forward arrow in the tool bar)
    - wait fo the kernel to restart and execute all cells (all executable cells should have numbers in front of them instead of a `[*]`) 
    - Check all executed cells for errors. If an exception is thrown, please check your code. Note that although the error might look cryptic, until now we never encounter that an exception was caused without a valid reason inside of the submitted code. A good way to check the code is to copy the solution into a new class in your favorite IDE and check
      - errors reported by the IDE
      - imports the IDE adds to your code which might be missing in your submission, here
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-b65aa29f897ebc4d, points: 0, grade: false, solution: true -->
```java
public int cut(int[] spaghetti, int maximumCuts) {
    // YOUR CODE HERE
    int position,j=1,maxSpaghetti=0;
    while(maximumCuts!=0 && (maximumCuts+j)*10>10) {
        position = searchElement(spaghetti,(maximumCuts+j)*10);
        if(position!=-1) {
            maximumCuts -= (spaghetti[position]/10-1);
            maxSpaghetti += spaghetti[position]/10;
        }
        j--;
    }
    
    int k=0;
    while(maximumCuts!=0 && k<spaghetti.length) {
        if(spaghetti[k]>10 && spaghetti[k] % 10!=0) {
            spaghetti[k] = spaghetti[k] - 10;
            maxSpaghetti++;
            maximumCuts--;
            if(spaghetti[k]<10){
                k++;
            }
        } else {
            k++;
        }
    }
       
    return maxSpaghetti;
}

public int searchElement(int[] spaghetti,int element) {
    for(int i=0;i<spaghetti.length;i++) {
        if(spaghetti[i]==element) {
            return i;
        }
    }
    return -1;
}
```

## EVALUATION AREA

In the following, we define some tests for evaluating your answer. The first test example can be used by you to test your implementation.

<!-- grade_id: cell-8bd565595479558a, points: 3, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

public void checkResult(int[] spaghettis, int maxCuts, int expectedResult) {
    try {
        long time = System.currentTimeMillis();
        int result = cut(spaghettis, maxCuts);
        time = System.currentTimeMillis() - time;
        Assertions.assertEquals(expectedResult, result);
        System.out.println("Test successful. Calculation took " + time + "ms");
    } catch (AssertionFailedError e) {
        throw e;
    } catch (Throwable e) {
        System.err.println("Your solution caused an unexpected error:");
        throw e;
    }
}
// The first parameter is the array of spaghettis,
// the second parameter is the maximum number of cuts and
// and the first number is the expected result
checkResult(new int[]{13,14,20}, 2, 3);
```

Output:
```
Test successful. Calculation took 0ms
```

<!-- grade_id: cell-41099f4980cf7022, points: 3, grade: true, solution: false -->
```java
// please ignore this cell :)
```

<!-- grade_id: cell-41099f4980cf7023, points: 3, grade: true, solution: false -->
```java
// please ignore this cell :)
```

<!-- grade_id: cell-41099f4980cf7024, points: 3, grade: true, solution: false -->
```java
// please ignore this cell :)
```

<!-- grade_id: cell-41099f4980cf7025, points: 3, grade: true, solution: false -->
```java
// please ignore this cell :)
```