# Data Science Suite
This is the assignment for the project group Data Science Suite. More information about the topics of this project group can be found at https://dice-research.org/teaching/DataScienceSuite4/.

The assignment comprises 3 files with one Java programming task per file. The files have no particular order. Please note that only working solutions will be granted points and that the style of programming (especially regarding the time and memory complexity) can be taken into account.

In case of severe issues or questions with one of the exercises you can send a mail to michael.roeder@uni-paderborn.de. Please note that only technical help can be granted.

Firstly, In which programming languages you feel confident? Please order the languages from most confident to least confident.

<!-- grade_id: cell-dd2b436fdf330ef1, points: 0, grade: true, solution: true -->
```java
/*
Name: Varun Nandkumar Golani
*/
```

# Print my HTML 
(15 points)

A program has created an HTML page as a tree of tags and their content. Every node is either a tag or contains some text. If it is a tag it might have one or more children. Nodes with text are not allowed to have any children. Every node has a reference to its parent (the root node has `null` as parent). You can see the implementation of the `Node` class below.

```java
/**
 * A simple class representing a markup language node in a tree.
 */
public static class Node {
    /**
     * The content of the node, i.e., either the name of the tag or the text the node is representing.
     */
    public final String content;
    /**
     * {@code true} if the node represents a tag, {@code false} if it represents text.
     */
    public final boolean isTag;
    /**
     * The parent of this node.
     */
    public final Node parent;
    /**
     * The children of this node (might be null!)
     */
    private List<Node> children = null;
    /**
     * Constructor creating the node and adding itself to the list of children of the given parent.
     * 
     * @param content the content of the node
     * @param isTag flag indicating whether the node is a tag
     * @param parent the parent of this node
     */
    public Node(String content, boolean isTag, Node parent) {
        this.content = content;
        this.isTag = isTag;
        this.parent = parent;
        if(parent != null) {
            parent.addChild(this);
        }
    }

    public List<Node> getChildren() {
        return children;
    }

    public void addChild(Node child) {
        if(children == null) {
            children = new ArrayList<>();
        }
        children.add(child);
    }
}
```

Your task is to extend the method `printHTML` shown below. I will receive the root node of a tree and should generate the correct String representation of the HTML page.

Imagine an easy HTML page where each line is a single `Node` instance and the indentation represents their hierarchy:
```
html
+-head
+-body
  +-p
    +-b
    | +-"Hello"
    +-"World!"
```

It can be seen that the page has an `html` element as root node. This node has `head` and `body` as children. While the `head` node comprises no children the `body` has a single `p` node as child. This `p` node contains a `b` element a the String `"World!"`. The `b` element has the String `"Hello"` as child.

This small tree can be summarized as a small HTML page with an empty head and a body which contains a paragraph with the text `"Hello World!"` in which the first word is bold. The output of your system should be a valid HTML String:
```xml
<html><head/><body><p><b>Hello</b>World!</p></body></html>
```

#### Hints
* For an easier comparison, whitespaces, tabs etc. will be removed.
* You should follow the XML rules for tags, i.e.,
  * `<tag>` is an opening tag
  * `</tag>` is a closing tag
  * `<tag/>` is an empty tag. Note that `<tag></tag>` will be accepted as well but **with a lower rating**

#### Notes

- You are free to use a different IDE to develop your solution. However, you have to copy the solution into this notebook to submit it.
- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Known issues
  - All global variables will be set to void after an import.
  - Missing spaces arround `%` (Modulo) can cause unexpected errors so please make sure that you have added spaces around every `%` character.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Make sure that
    - all necessary imports are listed at the beginning of your cell.
  - Run a final check of your solution by
    - click on _restart the kernel, then re-run the whole notebook_ (the fast forward arrow in the tool bar)
    - wait fo the kernel to restart and execute all cells (all executable cells should have numbers in front of them instead of a `[*]`) 
    - Check all executed cells for errors. If an exception is thrown, please check your code. Note that although the error might look cryptic, until now we never encounter that an exception was caused without a valid reason inside of the submitted code. A good way to check the code is to copy the solution into a new class in your favorite IDE and check
      - errors reported by the IDE
      - imports the IDE adds to your code which might be missing in your submission, here
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-b65aa29f897ebc4f, points: 0, grade: false, solution: true -->
```java
String htmlString = "";

//Reset the global variables used for each test case and calling helper function to print HTML
public String printHTML(Node root) {
    htmlString = "";
    printHTMLHelper(root);
    return htmlString;
}

public String printHTMLHelper(Node root) {
    // YOUR CODE HERE
    if(root == null){
        return "";
    } else {
        if(!root.isTag){
            htmlString = root.content;
        } else {
            htmlString += returnEmptyOrOpenCloseTags(root,"O");
            if(root.getChildren()!=null) {
                for (Node childNode : root.getChildren()) {
                    if(childNode.isTag) {
                        if(childNode.getChildren() == null){
                            htmlString += "<"+childNode.content+"/>";
                        } else {
                            printHTMLHelper(childNode);
                        }
                    } else {
                        htmlString += childNode.content;
                    }
                }
            }
            htmlString += returnEmptyOrOpenCloseTags(root,"C");
        }
        return htmlString;
    }
}

public String returnEmptyOrOpenCloseTags(Node node,String flag) {
    if(node.getChildren() == null && flag.equalsIgnoreCase("O")) {
        return "<"+node.content+"/>";
    } else if(node.getChildren() != null) {
        if(flag.equalsIgnoreCase("O")) {
            return "<"+node.content+">";
        } else if(flag.equalsIgnoreCase("C")) {
            return "</"+node.content+">"; 
        }
    }
    return "";
}
```

## EVALUATION AREA

In the following, we define some tests for evaluating your answer. The first test example can be used by you to test your implementation.

<!-- grade_id: cell-41099f4980cf7022, points: 1, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

// The example from the task description
Node html1 = new Node("html",true,null);
new Node("head",true,html1);
Node body1 = new Node("body",true,html1);
Node p1 = new Node("p",true,body1);
Node i1 = new Node("i",true,p1);
new Node("Hello",false,i1);
new Node("World!",false,p1);

long time = System.currentTimeMillis();
String result1 = printHTML(html1);
time = System.currentTimeMillis() - time;
result1 = result1.replaceAll("[\\s]+","");
Assertions.assertTrue(
"<html><head/><body><p><i>Hello</i>World!</p></body></html>".equals(result1) ||
"<html><head></head><body><p><i>Hello</i>World!</p></body></html>".equals(result1),
"\"" + result1 + "\" was not expected.");

System.out.println("Test was successful");
System.out.println("Needed " + time + "ms");
```

Output:
```
Test was successful
Needed 538ms
```

<!-- grade_id: cell-45a86fd3117e5a88, points: 5, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-45a86fd3117e5a8a, points: 2, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-45a86fd3117e5a8b, points: 2, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-b9f54341e2036963, points: 5, grade: true, solution: false -->
```java

```