# Nth prime
(15 points)

You should write a program which searches for the $N$-th prime number. The list of primes starts with:
```
2	3	5	7	11
```
Let 2 be the first prime number. It can be seen that 11 is the 5th prime. This means that if your program is called with the input $N=5$ it should return `11`.

The focus of this task is to implement a fast solution. The largest test will use $N=10000$. Please make sure that your program does not need more than a minute for this task to prevent a timeout of the evaluation process. **Note that a timeout yields 0 points!**. To get all points for this task, your program needs to have a very low runtime complexity and be able to solve the largest test in less than 200ms.

#### Notes

- You are free to use a different IDE to develop your solution. However, you have to copy the solution into this notebook to submit it.
- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Known issues
  - All global variables will be set to void after an import.
  - Missing spaces arround `%` (Modulo) can cause unexpected errors so please make sure that you have added spaces around every `%` character.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Make sure that
    - all necessary imports are listed at the beginning of your cell.
  - Run a final check of your solution by
    - click on _restart the kernel, then re-run the whole notebook_ (the fast forward arrow in the tool bar)
    - wait fo the kernel to restart and execute all cells (all executable cells should have numbers in front of them instead of a `[*]`) 
    - Check all executed cells for errors. If an exception is thrown, please check your code. Note that although the error might look cryptic, until now we never encounter that an exception was caused without a valid reason inside of the submitted code. A good way to check the code is to copy the solution into a new class in your favorite IDE and check
      - errors reported by the IDE
      - imports the IDE adds to your code which might be missing in your submission, here
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-b65aa29f897ebc4d, points: 0, grade: false, solution: true -->
```java
public int findNthPrime(int N) {
    int prime = 0,primeCount = 0,numberCount = 2;
    // YOUR CODE HERE
    while(primeCount<N) {
        if(isPrime(numberCount)) {
            primeCount++;
            prime = numberCount;
        }
        numberCount++;
    }
    return prime;
}

//Function to check if the number is prime or not
public boolean isPrime(int number) {
    int squareRoot = (int)Math.sqrt(number) + 1;
    
    if(number==1) {
        return false;
    }
    
    for(int i=2;i<squareRoot;i++) {
        if(number % i==0) {
            return false;
        }
    }
    return true;
}
```

## EVALUATION AREA

In the following, we define some tests for evaluating your answer. The first test example can be used by you to test your implementation.

<!-- grade_id: cell-41099f4980cf7020, points: 1, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

/**
 * Returns two numbers. The first is the Nth prime as calculated 
 * by the students solution. The second is the time needed for 
 * the calculation in ms.
 */
public long[] runExperiment(int N) {
    long[] results = new long[2];
    try {
        results[1] = System.currentTimeMillis();
        results[0] = findNthPrime(N);
        results[1] = System.currentTimeMillis() - results[1];
        return results;
    } catch (Throwable e) {
        System.err.println("Your solution caused an unexpected error:");
        throw e;
    }
}

public void checkResult(int expected, long calculated, long time) throws AssertionFailedError {
    Assertions.assertEquals(expected, calculated);
    System.out.println("Test successful. Calculation took " + time + "ms");
}

boolean testSuccessful = false;

long[] results = runExperiment(6);
// we expected 13
checkResult(13, results[0], results[1]);

results = runExperiment(100);
// we expected 541
checkResult(541, results[0], results[1]);

testSuccessful = true;
```

Output:
```
Test successful. Calculation took 0ms
Test successful. Calculation took 0ms
```

<!-- grade_id: cell-41099f4980cf7021, points: 1, grade: true, solution: false -->
```java
/**
 * We are preferring correct and fast solutions. Therefore, you get more points if your solution is fast.
 */
// make sure that the previous test didn't fail
Assertions.assertTrue(testSuccessful, "Previous test failed");
// check run time
Assertions.assertTrue(results[1] < 2, "Your solution took more time than expected.");
System.out.println("Test successful.");
```

Output:
```
Test successful.
```

<!-- grade_id: cell-41099f4980cf7022, points: 2, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-41099f4980cf7023, points: 2, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-41099f4980cf7024, points: 2, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-41099f4980cf7025, points: 2, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-41099f4980cf7026, points: 2, grade: true, solution: false -->
```java

```

<!-- grade_id: cell-41099f4980cf7027, points: 3, grade: true, solution: false -->
```java

```